[1596516275] SERVICE ALERT: esxi09;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.50 ms
[1596516339] SERVICE ALERT: esxi09;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1596524569] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596524689] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596524809] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596538199] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596538319] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;2;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596538439] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;HARD;3;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1596541121] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.37 ms
[1596541185] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1596541550] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.30 ms
[1596541614] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1596548266] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.42 ms
[1596548330] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms