[1544381511] SERVICE ALERT: router1;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544381514] HOST ALERT: router1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1544381662] SERVICE ALERT: router1;check-host-alive;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1544382324] HOST ALERT: router1;DOWN;HARD;10;(Host check timed out after 30.02 seconds)
[1544409028] HOST ALERT: router1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1544409235] SERVICE ALERT: router1;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1544422464] HOST ALERT: esxi05;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.55)
[1544422518] SERVICE ALERT: esxi05;PING;UNKNOWN;HARD;1;check_ping: %s: Warning threshold must be integer or percentage!
[1544422558] SERVICE ALERT: esxi04;PING;UNKNOWN;SOFT;1;check_ping: %s: Warning threshold must be integer or percentage!
[1544422618] SERVICE ALERT: esxi04;PING;UNKNOWN;SOFT;2;check_ping: %s: Warning threshold must be integer or percentage!
[1544422678] SERVICE ALERT: esxi04;PING;UNKNOWN;SOFT;3;check_ping: %s: Warning threshold must be integer or percentage!
[1544422738] SERVICE ALERT: esxi04;PING;UNKNOWN;HARD;4;check_ping: %s: Warning threshold must be integer or percentage!
[1544422793] SERVICE ALERT: esxi05;ESXi Hardware Check;UNKNOWN;HARD;1;UNKNOWN: (0, 'Socket error: [Errno 113] No route to host')
[1544422979] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544423031] HOST ALERT: esxi05;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.55)
[1544423099] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;2;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544423112] SERVICE ALERT: esxi06;PING;UNKNOWN;SOFT;1;check_ping: %s: Warning threshold must be integer or percentage!
[1544423172] SERVICE ALERT: esxi06;PING;UNKNOWN;SOFT;2;check_ping: %s: Warning threshold must be integer or percentage!
[1544423230] SERVICE ALERT: esxi06;ESXi Haradware Check;OK;HARD;1;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERW Chassis S/N: FCH2216V0X8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1544423232] SERVICE ALERT: esxi06;PING;UNKNOWN;SOFT;3;check_ping: %s: Warning threshold must be integer or percentage!
[1544423296] SERVICE ALERT: esxi06;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544423342] SERVICE ALERT: esxi04;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544423481] SERVICE ALERT: esxi05;PING;CRITICAL;HARD;4;CRITICAL - Host Unreachable (10.5.5.55)
[1544424197] SERVICE ALERT: esxi06;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.73 ms
[1544424261] SERVICE ALERT: esxi06;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1544424327] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544424447] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544424567] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544425447] SERVICE ALERT: esxi04;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.50 ms
[1544425511] SERVICE ALERT: esxi04;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1544426651] SERVICE ALERT: esxi01;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.48 ms
[1544426715] SERVICE ALERT: esxi01;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1544427300] SERVICE ALERT: esxi09;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544427316] SERVICE ALERT: esxi01;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.50 ms
[1544427380] SERVICE ALERT: esxi01;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544427420] SERVICE ALERT: esxi09;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544427540] SERVICE ALERT: esxi09;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN: (0, 'Socket error: [Errno 111] Connection refused')
[1544431561] HOST ALERT: router1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1544431625] HOST ALERT: router1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.71 ms