[1741700533] SERVICE ALERT: esxi07;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.5.5.57)
[1741700541] SERVICE ALERT: esxi06;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1741700541] SERVICE ALERT: esxi05;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1741700543] HOST ALERT: esxi06;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.56)
[1741700543] SERVICE ALERT: esxi01;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EQH Chassis S/N: FCH2216V0WN  System BIOS: C220M4.4.1.2c.0.0202211901 2021-02-02
[1741700543] HOST ALERT: esxi07;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.57)
[1741700547] HOST ALERT: esxi05;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.55)
[1741700595] SERVICE ALERT: esxi07;PING;CRITICAL;HARD;4;CRITICAL - Host Unreachable (10.5.5.57)
[1741700604] SERVICE ALERT: esxi06;PING;CRITICAL;HARD;4;CRITICAL - Host Unreachable (10.5.5.56)
[1741700604] HOST ALERT: esxi06;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.5.5.56)
[1741700604] SERVICE ALERT: esxi05;PING;CRITICAL;HARD;4;CRITICAL - Host Unreachable (10.5.5.55)
[1741700606] HOST ALERT: esxi07;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.5.5.57)
[1741700611] HOST ALERT: esxi05;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.5.5.55)
[1741700667] HOST ALERT: esxi06;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.5.5.56)
[1741700670] HOST ALERT: esxi07;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.5.5.57)
[1741700676] HOST ALERT: esxi05;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 1.64 ms
[1741700731] HOST ALERT: esxi06;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1741700734] HOST ALERT: esxi07;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1741700897] SERVICE ALERT: esxi07;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1741700905] SERVICE ALERT: esxi06;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1741700905] SERVICE ALERT: esxi05;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.38 ms