[1539670543] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.003 second response time
[1539670603] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1539670663] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.002 second response time
[1539670723] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1539683300] SERVICE ALERT: localhost.localdomain;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5168 bytes in 0.011 second response time
[1539683360] SERVICE ALERT: localhost.localdomain;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5168 bytes in 0.004 second response time
[1539683420] SERVICE ALERT: localhost.localdomain;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5168 bytes in 0.005 second response time
[1539683480] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;HARD;4;connect to address 10.5.5.15 and port 80: Connection refused
[1539683780] SERVICE ALERT: localhost.localdomain;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5168 bytes in 0.006 second response time
[1539684380] SERVICE ALERT: localhost.localdomain;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 285 bytes in 0.004 second response time
[1539684540] SERVICE ALERT: localhost.localdomain;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539684543] HOST ALERT: localhost.localdomain;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539684601] SERVICE ALERT: localhost.localdomain;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539684612] SERVICE ALERT: localhost.localdomain;SSH;CRITICAL;HARD;1;connect to address 10.5.5.15 and port 22: No route to host
[1539684682] SERVICE FLAPPING ALERT: localhost.localdomain;HTTP;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1539684682] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;HARD;1;connect to address 10.5.5.15 and port 80: No route to host
[1539684931] SERVICE FLAPPING ALERT: localhost.localdomain;HTTP;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1539684941] HOST ALERT: localhost.localdomain;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.15)
[1539684961] SERVICE FLAPPING ALERT: localhost.localdomain;HTTP;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1539685029] HOST ALERT: localhost.localdomain;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1539685042] SERVICE ALERT: localhost.localdomain;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 285 bytes in 0.005 second response time
[1539685204] SERVICE ALERT: localhost.localdomain;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1539685272] SERVICE ALERT: localhost.localdomain;SSH;OK;HARD;1;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1539685642] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;SOFT;1;connect to address 10.5.5.15 and port 80: Connection refused
[1539685712] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1539685775] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;SOFT;3;connect to address 10.5.5.15 and port 80: No route to host
[1539685807] SERVICE ALERT: localhost.localdomain;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539685810] HOST ALERT: localhost.localdomain;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539685838] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;HARD;3;connect to address 10.5.5.15 and port 80: No route to host
[1539685870] SERVICE ALERT: localhost.localdomain;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539685874] SERVICE ALERT: localhost.localdomain;SSH;CRITICAL;HARD;1;connect to address 10.5.5.15 and port 22: No route to host
[1539686228] HOST ALERT: localhost.localdomain;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.15)
[1539688859] SERVICE FLAPPING ALERT: localhost.localdomain;HTTP;STARTED; Service appears to have started flapping (32.5% change >= 20.0% threshold)
[1539689144] HOST ALERT: localhost.localdomain;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1539689290] SERVICE ALERT: localhost.localdomain;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 285 bytes in 0.005 second response time
[1539689331] SERVICE FLAPPING ALERT: localhost.localdomain;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1539689331] SERVICE ALERT: localhost.localdomain;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1539689402] SERVICE FLAPPING ALERT: localhost.localdomain;SSH;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1539689402] SERVICE ALERT: localhost.localdomain;SSH;OK;HARD;1;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1539691390] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;SOFT;1;connect to address 10.5.5.15 and port 80: Connection refused
[1539691460] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1539691505] SERVICE ALERT: localhost.localdomain;SSH;CRITICAL;SOFT;1;connect to address 10.5.5.15 and port 22: No route to host
[1539691508] HOST ALERT: localhost.localdomain;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539691523] SERVICE ALERT: localhost.localdomain;HTTP;CRITICAL;HARD;2;connect to address 10.5.5.15 and port 80: No route to host
[1539691568] SERVICE ALERT: localhost.localdomain;SSH;CRITICAL;HARD;1;connect to address 10.5.5.15 and port 22: No route to host
[1539691734] SERVICE ALERT: localhost.localdomain;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.15)
[1539691955] HOST ALERT: localhost.localdomain;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.15)
[1539695095] SERVICE FLAPPING ALERT: localhost.localdomain;PING;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)