[1601321821] SERVICE ALERT: esxi05;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.90 ms
[1601321885] SERVICE ALERT: esxi05;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1601350804] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1601350807] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1601350808] HOST ALERT: Director_Bldg;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1601350811] HOST ALERT: Director_Resident;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1601350898] HOST ALERT: Director_Bldg;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1601350901] HOST ALERT: Director_Resident;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1601350955] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1601350958] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1601350988] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.02 seconds)
[1601350991] HOST ALERT: Director_Resident;DOWN;SOFT;3;(Host check timed out after 30.02 seconds)
[1601351056] HOST ALERT: Director_Resident;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 0.79 ms
[1601351078] HOST ALERT: Director_Bldg;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1601351168] HOST ALERT: Director_Bldg;DOWN;SOFT;5;(Host check timed out after 30.03 seconds)
[1601351258] HOST ALERT: Director_Bldg;DOWN;SOFT;6;(Host check timed out after 30.03 seconds)
[1601351348] HOST ALERT: Director_Bldg;DOWN;SOFT;7;(Host check timed out after 30.02 seconds)
[1601351438] HOST ALERT: Director_Bldg;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1601351528] HOST ALERT: Director_Bldg;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1601351531] SERVICE ALERT: Director_Resident;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1601351618] HOST ALERT: Director_Bldg;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1601351922] HOST ALERT: Director_Bldg;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1601352128] SERVICE ALERT: Director_Bldg;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1601365300] SERVICE ALERT: esxi06;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.46 ms
[1601365364] SERVICE ALERT: esxi06;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms