[1651262650] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.48 ms
[1651262714] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1651296335] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1651296335] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1651296337] HOST ALERT: Director_Bldg;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.3.1.10)
[1651296338] HOST ALERT: Director_Resident;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.1.1.26)
[1651296400] HOST ALERT: Director_Bldg;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.3.1.10)
[1651296401] HOST ALERT: Director_Resident;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.1.1.26)
[1651296458] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.1.1.26)
[1651296458] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.3.1.10)
[1651296464] HOST ALERT: Director_Resident;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.1.1.26)
[1651296490] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.02 seconds)
[1651296527] HOST ALERT: Director_Resident;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.1.1.26)
[1651296551] HOST ALERT: Director_Bldg;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.3.1.10)
[1651296591] HOST ALERT: Director_Resident;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.1.1.26)
[1651296614] HOST ALERT: Director_Bldg;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.3.1.10)
[1651296654] HOST ALERT: Director_Resident;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.1.1.26)
[1651296677] HOST ALERT: Director_Bldg;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.3.1.10)
[1651296717] HOST ALERT: Director_Resident;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.1.1.26)
[1651296740] HOST ALERT: Director_Bldg;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.3.1.10)
[1651296780] HOST ALERT: Director_Resident;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.1.1.26)
[1651296803] HOST ALERT: Director_Bldg;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.3.1.10)
[1651296843] HOST ALERT: Director_Resident;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.1.1.26)
[1651296866] HOST ALERT: Director_Bldg;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.3.1.10)
[1651296906] HOST ALERT: Director_Resident;DOWN;HARD;10;CRITICAL - Host Unreachable (10.1.1.26)
[1651296929] HOST ALERT: Director_Bldg;DOWN;HARD;10;CRITICAL - Host Unreachable (10.3.1.10)
[1651307133] HOST ALERT: Director_Bldg;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1651307259] SERVICE ALERT: Director_Bldg;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1651307259] SERVICE ALERT: Director_Resident;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1651307263] HOST ALERT: Director_Resident;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1651323079] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.46 ms
[1651323143] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms