[1555704329] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.46 ms
[1555704393] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1555709258] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.49 ms
[1555709322] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1555709387] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.40 ms
[1555709451] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1555711316] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.49 ms
[1555711380] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1555715645] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.45 ms
[1555715709] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1555723574] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.49 ms
[1555723638] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1555733003] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.42 ms
[1555733067] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1555739937] SERVICE ALERT: esxi03;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ESB Chassis S/N: FCH2216V0TW  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740065] SERVICE ALERT: esxi03;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ESB Chassis S/N: FCH2216V0TW  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740139] SERVICE ALERT: esxi04;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QKB Chassis S/N: FCH2216V0YZ  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740143] SERVICE ALERT: esxi01;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EQH Chassis S/N: FCH2216V0WN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740194] SERVICE ALERT: esxi03;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ESB Chassis S/N: FCH2216V0TW  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740233] SERVICE ALERT: esxi02;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157E1D Chassis S/N: FCH2216V0Y6  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740268] SERVICE ALERT: esxi04;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QKB Chassis S/N: FCH2216V0YZ  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740272] SERVICE ALERT: esxi01;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EQH Chassis S/N: FCH2216V0WN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740361] SERVICE ALERT: esxi02;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157E1D Chassis S/N: FCH2216V0Y6  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740397] SERVICE ALERT: esxi04;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QKB Chassis S/N: FCH2216V0YZ  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740400] SERVICE ALERT: esxi01;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EQH Chassis S/N: FCH2216V0WN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740491] SERVICE ALERT: esxi02;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157E1D Chassis S/N: FCH2216V0Y6  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740544] SERVICE ALERT: esxi05;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EPD Chassis S/N: FCH2216V0Z0  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740672] SERVICE ALERT: esxi05;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EPD Chassis S/N: FCH2216V0Z0  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740802] SERVICE ALERT: esxi05;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EPD Chassis S/N: FCH2216V0Z0  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740835] SERVICE ALERT: esxi06;ESXi Haradware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERW Chassis S/N: FCH2216V0X8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740953] SERVICE ALERT: esxi07;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EV7 Chassis S/N: FCH2216V0Y8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555740964] SERVICE ALERT: esxi06;ESXi Haradware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERW Chassis S/N: FCH2216V0X8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555741082] SERVICE ALERT: esxi07;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EV7 Chassis S/N: FCH2216V0Y8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555741093] SERVICE ALERT: esxi06;ESXi Haradware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERW Chassis S/N: FCH2216V0X8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555741210] SERVICE ALERT: esxi07;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EV7 Chassis S/N: FCH2216V0Y8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555741787] SERVICE ALERT: esxi09;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERY Chassis S/N: FCH2216V0YN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555741915] SERVICE ALERT: esxi09;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERY Chassis S/N: FCH2216V0YN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555742043] SERVICE ALERT: esxi09;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERY Chassis S/N: FCH2216V0YN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555742053] SERVICE ALERT: esxi08;ESXi Hardware Check;WARNING;SOFT;1;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QJN Chassis S/N: FCH2216V0YV  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555742181] SERVICE ALERT: esxi08;ESXi Hardware Check;WARNING;SOFT;2;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 40  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QJN Chassis S/N: FCH2216V0YV  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555742309] SERVICE ALERT: esxi08;ESXi Hardware Check;WARNING;HARD;3;WARNING : Front Panel Board 1 FP_TEMP_SENSOR: 41  - Server:  Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QJN Chassis S/N: FCH2216V0YV  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744006] SERVICE ALERT: esxi04;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QKB Chassis S/N: FCH2216V0YZ  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744101] SERVICE ALERT: esxi02;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157E1D Chassis S/N: FCH2216V0Y6  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744103] SERVICE ALERT: esxi06;ESXi Haradware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERW Chassis S/N: FCH2216V0X8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744121] SERVICE ALERT: esxi08;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QJN Chassis S/N: FCH2216V0YV  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744219] SERVICE ALERT: esxi07;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EV7 Chassis S/N: FCH2216V0Y8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744403] SERVICE ALERT: esxi03;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ESB Chassis S/N: FCH2216V0TW  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744412] SERVICE ALERT: esxi05;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EPD Chassis S/N: FCH2216V0Z0  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744451] SERVICE ALERT: esxi09;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERY Chassis S/N: FCH2216V0YN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555744609] SERVICE ALERT: esxi01;ESXi Hardware Check;OK;HARD;3;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EQH Chassis S/N: FCH2216V0WN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1555746632] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.60 ms
[1555746696] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1555747361] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.36 ms
[1555747425] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1555755890] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.36 ms
[1555755954] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1555762619] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 3.01 ms
[1555762683] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1555763048] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.43 ms
[1555763112] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1555764077] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.39 ms
[1555764141] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1555767506] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.40 ms
[1555767570] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1555772735] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.45 ms
[1555772799] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1555773164] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.45 ms
[1555773228] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1555775393] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.40 ms
[1555775457] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1555775522] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.38 ms
[1555775586] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1555777151] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.41 ms
[1555777215] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1555777880] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.44 ms
[1555777944] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms