[1575485602] SERVICE ALERT: esxi01;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.50 ms
[1575485666] SERVICE ALERT: esxi01;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1575495673] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.61 ms
[1575495737] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1575509173] SERVICE ALERT: esxi07;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.47 ms
[1575509237] SERVICE ALERT: esxi07;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1575509772] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.73 ms
[1575509836] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1575511702] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.45 ms
[1575511766] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1575532469] SERVICE ALERT: esxi01;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575532471] SERVICE ALERT: esxi02;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575532621] SERVICE ALERT: esxi01;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575532623] SERVICE ALERT: esxi02;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575532772] SERVICE ALERT: esxi01;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN:
[1575532774] SERVICE ALERT: esxi02;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN:
[1575533098] SERVICE ALERT: esxi08;ESXi Hardware Check;CRITICAL;SOFT;1;(Service check timed out after 60.02 seconds)
[1575533100] SERVICE ALERT: esxi09;ESXi Hardware Check;CRITICAL;SOFT;1;(Service check timed out after 60.02 seconds)
[1575533278] SERVICE ALERT: esxi08;ESXi Hardware Check;CRITICAL;SOFT;2;(Service check timed out after 60.06 seconds)
[1575533280] SERVICE ALERT: esxi09;ESXi Hardware Check;CRITICAL;SOFT;2;(Service check timed out after 60.06 seconds)
[1575533458] SERVICE ALERT: esxi08;ESXi Hardware Check;CRITICAL;HARD;3;(Service check timed out after 60.02 seconds)
[1575533460] SERVICE ALERT: esxi09;ESXi Hardware Check;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1575545668] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575545669] SERVICE ALERT: esxi03;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575545669] SERVICE ALERT: esxi05;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575545671] SERVICE ALERT: esxi04;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN:
[1575545819] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575545820] SERVICE ALERT: esxi03;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575545820] SERVICE ALERT: esxi05;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575545821] SERVICE ALERT: esxi04;ESXi Hardware Check;UNKNOWN;SOFT;2;UNKNOWN:
[1575545971] SERVICE ALERT: esxi06;ESXi Haradware Check;UNKNOWN;HARD;3;UNKNOWN:
[1575545971] SERVICE ALERT: esxi03;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN:
[1575545972] SERVICE ALERT: esxi05;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN:
[1575545973] SERVICE ALERT: esxi04;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN: