[1603393460] SERVICE ALERT: esxi05;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.43 ms
[1603393524] SERVICE ALERT: esxi05;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1603436601] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1603436631] HOST ALERT: Director_Resident;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1603436721] HOST ALERT: Director_Resident;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1603436752] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1603436786] HOST ALERT: Director_Resident;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 0.70 ms
[1603437325] SERVICE ALERT: Director_Resident;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1603453285] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1603453311] HOST ALERT: Director_Resident;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1603453315] HOST ALERT: Director_Bldg;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1603453401] HOST ALERT: Director_Resident;DOWN;SOFT;2;(Host check timed out after 30.02 seconds)
[1603453405] HOST ALERT: Director_Bldg;DOWN;SOFT;2;(Host check timed out after 30.02 seconds)
[1603453436] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1603453491] HOST ALERT: Director_Resident;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1603453495] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1603453556] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1603453581] HOST ALERT: Director_Resident;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1603453585] HOST ALERT: Director_Bldg;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1603453671] HOST ALERT: Director_Resident;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1603453675] HOST ALERT: Director_Bldg;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1603453761] HOST ALERT: Director_Resident;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1603453765] HOST ALERT: Director_Bldg;DOWN;SOFT;6;(Host check timed out after 30.04 seconds)
[1603453851] HOST ALERT: Director_Resident;DOWN;SOFT;7;(Host check timed out after 30.03 seconds)
[1603453855] HOST ALERT: Director_Bldg;DOWN;SOFT;7;(Host check timed out after 30.03 seconds)
[1603453941] HOST ALERT: Director_Resident;DOWN;SOFT;8;(Host check timed out after 30.02 seconds)
[1603453945] HOST ALERT: Director_Bldg;DOWN;SOFT;8;(Host check timed out after 30.03 seconds)
[1603454031] HOST ALERT: Director_Resident;DOWN;SOFT;9;(Host check timed out after 30.02 seconds)
[1603454035] HOST ALERT: Director_Bldg;DOWN;SOFT;9;(Host check timed out after 30.03 seconds)
[1603454121] HOST ALERT: Director_Resident;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1603454125] HOST ALERT: Director_Bldg;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1603455626] HOST ALERT: Director_Resident;UP;HARD;1;PING OK - Packet loss = 20%, RTA = 0.75 ms
[1603455629] HOST ALERT: Director_Bldg;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1603455809] SERVICE ALERT: Director_Bldg;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1603455929] SERVICE ALERT: Director_Resident;check-host-alive;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1603461403] SERVICE ALERT: esxi04;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.52 ms
[1603461467] SERVICE ALERT: esxi04;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms