[1584647215] SERVICE ALERT: esxi04;ESXi Hardware Check;UNKNOWN;HARD;3;UNKNOWN:
[1584690000] HOST ALERT: Director_Bldg;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1584690090] HOST ALERT: Director_Bldg;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1584690180] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.03 seconds)
[1584690270] HOST ALERT: Director_Bldg;DOWN;SOFT;4;(Host check timed out after 30.03 seconds)
[1584690294] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;1;CRITICAL - Plugin timed out
[1584690360] HOST ALERT: Director_Bldg;DOWN;SOFT;5;(Host check timed out after 30.03 seconds)
[1584690450] HOST ALERT: Director_Bldg;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1584690540] HOST ALERT: Director_Bldg;DOWN;SOFT;7;(Host check timed out after 30.03 seconds)
[1584690630] HOST ALERT: Director_Bldg;DOWN;SOFT;8;(Host check timed out after 30.03 seconds)
[1584690720] HOST ALERT: Director_Bldg;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1584690810] HOST ALERT: Director_Bldg;DOWN;HARD;10;(Host check timed out after 30.03 seconds)
[1584694907] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.31 ms
[1584694971] SERVICE ALERT: esxi03;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1584713314] HOST ALERT: Director_Bldg;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1584713667] SERVICE ALERT: Director_Bldg;check-host-alive;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.58 ms