[1544681613] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.59 ms
[1544681677] SERVICE ALERT: esxi02;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1544686516] SERVICE ALERT: esxi03;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.46 ms
[1544686580] SERVICE ALERT: esxi03;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1544698514] SERVICE ALERT: esxi08;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544698516] HOST ALERT: esxi08;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.58)
[1544698519] SERVICE ALERT: esxi08;ESXi Hardware Check;UNKNOWN;HARD;1;UNKNOWN: (0, 'Socket error: [Errno 113] No route to host')
[1544698532] SERVICE ALERT: esxi09;ESXi Hardware Check;UNKNOWN;SOFT;1;UNKNOWN: (0, 'Socket error: [Errno 113] No route to host')
[1544698535] HOST ALERT: esxi09;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.59)
[1544698565] SERVICE ALERT: router1;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544698568] HOST ALERT: router1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.3.1.10)
[1544698577] SERVICE ALERT: esxi08;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.58)
[1544698585] SERVICE ALERT: router3;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.1.1.26)
[1544698588] HOST ALERT: router3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.1.1.26)
[1544698626] SERVICE ALERT: esxi07;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544698629] HOST ALERT: esxi07;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.57)
[1544698647] SERVICE ALERT: esxi07;ESXi Hardware Check;UNKNOWN;HARD;1;UNKNOWN: (0, 'Socket error: [Errno 113] No route to host')
[1544698649] HOST ALERT: vmwaremgmtswitch;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.1.1.39)
[1544698658] SERVICE ALERT: esxi09;ESXi Hardware Check;UNKNOWN;HARD;1;UNKNOWN: (0, 'Socket error: [Errno 113] No route to host')
[1544698688] SERVICE ALERT: router1;check-host-alive;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.3.1.10)
[1544698689] SERVICE ALERT: esxi07;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.57)
[1544698701] HOST ALERT: router2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.3.1.25)
[1544698707] SERVICE ALERT: router3;check-host-alive;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.1.1.26)
[1544698710] SERVICE ALERT: esxi09;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.5.5.59)
[1544698723] HOST ALERT: vmware10gswitch;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.1.1.40)
[1544698849] SERVICE ALERT: vmwaremgmtswitch;check-host-alive;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.1.1.39)
[1544698849] SERVICE ALERT: vmware10gswitch;check-host-alive;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.1.1.40)
[1544699038] SERVICE ALERT: router2;check-host-alive;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.3.1.25)
[1544699081] HOST ALERT: esxi08;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.58)
[1544699099] HOST ALERT: esxi09;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.59)
[1544699132] HOST ALERT: router1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.3.1.10)
[1544699150] HOST ALERT: router3;DOWN;HARD;10;CRITICAL - Host Unreachable (10.1.1.26)
[1544699196] HOST ALERT: esxi07;DOWN;HARD;10;CRITICAL - Host Unreachable (10.5.5.57)
[1544699213] HOST ALERT: vmwaremgmtswitch;DOWN;HARD;10;CRITICAL - Host Unreachable (10.1.1.39)
[1544699266] HOST ALERT: router2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.3.1.25)
[1544699287] HOST ALERT: vmware10gswitch;DOWN;HARD;10;CRITICAL - Host Unreachable (10.1.1.40)
[1544700425] SERVICE ALERT: esxi01;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544700447] SERVICE ALERT: esxi08;ESXi Hardware Check;OK;HARD;1;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22137QJN Chassis S/N: FCH2216V0YV  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1544700451] HOST ALERT: esxi08;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1544700461] SERVICE ALERT: esxi09;ESXi Hardware Check;OK;HARD;1;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157ERY Chassis S/N: FCH2216V0YN  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1544700465] HOST ALERT: esxi09;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544700470] HOST ALERT: router2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1544700489] SERVICE ALERT: router1;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1544700490] SERVICE ALERT: esxi01;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544700490] SERVICE ALERT: esxi07;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1544700491] HOST ALERT: vmware10gswitch;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.31 ms
[1544700493] HOST ALERT: router1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1544700494] HOST ALERT: esxi07;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1544700509] SERVICE ALERT: router3;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1544700513] HOST ALERT: router3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1544700574] SERVICE ALERT: esxi09;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1544700577] SERVICE ALERT: esxi07;ESXi Hardware Check;OK;HARD;1;OK - Server: Cisco Systems Inc UCSC-C220-M4S s/n: FCH22157EV7 Chassis S/N: FCH2216V0Y8  System BIOS: C220M4.3.0.4a.0.0226182253 2018-02-26
[1544700678] SERVICE ALERT: esxi08;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1544700717] HOST ALERT: vmwaremgmtswitch;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 4.23 ms
[1544700773] SERVICE ALERT: vmwaremgmtswitch;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.84 ms
[1544700773] SERVICE ALERT: vmware10gswitch;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1544700962] SERVICE ALERT: router2;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1544701585] SERVICE ALERT: esxi08;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544701588] HOST ALERT: esxi08;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.5.5.58)
[1544701589] SERVICE ALERT: router2;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1544701592] HOST ALERT: router2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.3.1.25)
[1544701616] HOST ALERT: vmwaremgmtswitch;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.1.1.39)
[1544701649] SERVICE ALERT: esxi08;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1544701652] HOST ALERT: esxi08;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1544701656] HOST ALERT: router2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1544701680] HOST ALERT: vmwaremgmtswitch;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1544701713] SERVICE ALERT: router2;check-host-alive;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1544708678] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.35 ms
[1544708742] SERVICE ALERT: esxi02;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1544720890] SERVICE ALERT: esxi01;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.55 ms
[1544720954] SERVICE ALERT: esxi01;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1544723675] SERVICE ALERT: esxi09;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.51 ms
[1544723739] SERVICE ALERT: esxi09;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms