[1607745928] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1607745932] HOST ALERT: Director_Resident;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1607745949] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1607745952] HOST ALERT: Director_Bldg;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1607746022] HOST ALERT: Director_Resident;DOWN;SOFT;2;(Host check timed out after 30.03 seconds)
[1607746042] HOST ALERT: Director_Bldg;DOWN;SOFT;2;(Host check timed out after 30.03 seconds)
[1607746079] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1607746100] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1607746112] HOST ALERT: Director_Resident;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1607746132] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1607746177] HOST ALERT: Director_Resident;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 1.24 ms
[1607746196] HOST ALERT: Director_Bldg;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1607746652] SERVICE ALERT: Director_Resident;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1607746673] SERVICE ALERT: Director_Bldg;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.67 ms