[1575215040] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 5.11, 1.91, 0.70
[1575215044] SERVICE ALERT: esxi02;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.33 ms
[1575215069] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1575215069] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1575215086] HOST ALERT: Director_Bldg;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1575215099] HOST ALERT: Director_Resident;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1575215100] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 6.14, 2.99, 1.16
[1575215108] SERVICE ALERT: esxi02;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1575215160] SERVICE ALERT: localhost;Current Load;OK;SOFT;3;OK - load average: 4.58, 3.14, 1.32
[1575215176] HOST ALERT: Director_Bldg;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1575215189] HOST ALERT: Director_Resident;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1575215220] SERVICE ALERT: Director_Resident;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1575215220] SERVICE ALERT: Director_Bldg;check-host-alive;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1575215266] HOST ALERT: Director_Bldg;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1575215279] HOST ALERT: Director_Resident;DOWN;SOFT;3;(Host check timed out after 30.03 seconds)
[1575215356] HOST ALERT: Director_Bldg;DOWN;SOFT;4;(Host check timed out after 30.02 seconds)
[1575215369] HOST ALERT: Director_Resident;DOWN;SOFT;4;(Host check timed out after 30.02 seconds)
[1575215446] HOST ALERT: Director_Bldg;DOWN;SOFT;5;(Host check timed out after 30.02 seconds)
[1575215459] HOST ALERT: Director_Resident;DOWN;SOFT;5;(Host check timed out after 30.03 seconds)
[1575215536] HOST ALERT: Director_Bldg;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1575215549] HOST ALERT: Director_Resident;DOWN;SOFT;6;(Host check timed out after 30.03 seconds)
[1575215626] HOST ALERT: Director_Bldg;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1575215639] HOST ALERT: Director_Resident;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1575215716] HOST ALERT: Director_Bldg;DOWN;SOFT;8;(Host check timed out after 30.03 seconds)
[1575215729] HOST ALERT: Director_Resident;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1575215806] HOST ALERT: Director_Bldg;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1575215819] HOST ALERT: Director_Resident;DOWN;SOFT;9;(Host check timed out after 30.03 seconds)
[1575215896] HOST ALERT: Director_Bldg;DOWN;HARD;10;(Host check timed out after 30.02 seconds)
[1575215909] HOST ALERT: Director_Resident;DOWN;HARD;10;(Host check timed out after 30.03 seconds)
[1575216843] SERVICE ALERT: esxi06;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 3.45 ms
[1575216907] SERVICE ALERT: esxi06;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1575217400] HOST ALERT: Director_Bldg;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1575217414] HOST ALERT: Director_Resident;UP;HARD;1;PING OK - Packet loss = 40%, RTA = 0.74 ms
[1575217593] SERVICE ALERT: Director_Bldg;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1575217593] SERVICE ALERT: Director_Resident;check-host-alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1575217743] SERVICE ALERT: esxi04;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 40%, RTA = 0.44 ms
[1575217807] SERVICE ALERT: esxi04;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1575218343] SERVICE ALERT: esxi08;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 0.69 ms
[1575218407] SERVICE ALERT: esxi08;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms